---
title       : представление данных
description : описательная статистика
--- type:NormalExercise lang:r xp:100 skills:1 key:3f38ea43bc
## Упражнение 1


*** =instructions
Вычислите среднее
*** =hint
используйте функцию mean 
*** =pre_exercise_code
```{a <- binomial(1, 2)
    b <- 2
}

```

*** =sample_code
```{c <- mean(a, b)}

```

*** =solution
```{c <- mean(a, b)}

```

*** =sct
```{mean(a, b)}

```
