---
title       : представление данных
description : описательная статистика
--- type:NormalExercise lang:r xp:100 skills:1 key:3f38ea43bc
## Упражнение 1


*** =instructions
Вычислите среднее
*** =hint
используйте функцию mean 
*** =pre_exercise_code
```{r}

```

*** =sample_code
```{r}

```

*** =solution
```{r}

```

*** =sct
```{r}

```
