---
title       : представление данных
description : описательная статистика
--- type:NormalExercise lang:r xp:100 skills:1 key:3f38ea43bc
## Упражнение 1


*** =instructions
Вычислите среднее
*** =hint
используйте функцию mean 
*** =pre_exercise_code
```{r}
a <- 1

```

*** =sample_code
```{r}
a <- 1

```

*** =solution
```{r}
a <- 5

```

*** =sct
```{r}
test_error()

```
